Emotional Vocalizations Alter Behaviors and Neurochemical Release into the Amygdala

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

The manuscript addresses a fundamental question about how different types of communication signals differentially affect brain state and neurochemistry. In addition, their manuscript highlights the various processes that modulate brain responses to communication signals, including prior experience, sex, and hormonal status. Overall, the manuscript is well-written and the research is appropriately contextualized.

That being said, it remains important for the authors to think more about their analytical approaches. In particular, the effect of normalization and the explicit outlining and interpretations of statistical models. As mentioned in the original review, the normalization of neurochemical data seems unnecessary given the repeated-measures design of their analysis and by normalizing all data to the baseline data and including this baseline data in the repeated measures analysis, one artificially creates a baseline period with minimal variation that dramatically differs in variance from other periods (akin to heteroscedasticity). If the authors want to analyze how a stimulus changes neurochemical concentrations, they could analyze the raw data but depict normalized data in their figures (similar to other papers). Or they could analyze group differences in the normalized data of the two stimulus periods (i.e., excluding the baseline period used for normalization).

We appreciate the reviewer’s point on the difference in variance caused by including the 100% baseline values in the analysis. After consulting with our statistician, we chose the latter of the two approaches suggested by the reviewer. Specifically, we reran the analysis to exclude the baseline and focus only on the playback windows and the group differences. The text in the results, the significance signs in the figures, and the discussion are corrected accordingly. Despite these changes, our major conclusions remains as before.

We also followed this reviewer’s suggestions to clarify the statistical model in studying the experience effect. After further consultation with our statistician, we reran the analysis on experience effect, including all the groups of EXP and INEXP animals together. We have corrected text in the figure captions, results, discussion, and data analysis sections of the manuscript related to the effect of experience and its interactions. This has not changed the conclusion made related to the experience effect in the dataset.

It would also be useful for the authors to provide further discussion of the potential contributions of different types of experiences (mating vs. restraint) to the change in behavior and neurochemical responses to the vocalization playbacks and to try to disentangle sensory and motor contributions to neurochemical changes.

We have acknowledged in the Discussion that previous studies suggest that the effect of experience involving stress could be generalized. We believe that this is an important area of future research. Our Discussion acknowledges that the relationship between sensory and motor contributions to neurochemical changes remains an area of interest. We further point out that the time resolution of microdialysis data renders the suggested discussion highly speculative. We plan to use other methods to assess this in future experiments.

Reviewer #3 (Public Review):

The work by Ghasemahmad et al. has the potential to significantly advance our understanding of how neuromodulators provide internal-state signals to the basolateral amygdala (BLA) while an animal listens to social vocalizations.

Ghasemahmad et al. made changes to the manuscript that have significantly improved the work. In particular, the transparency in showing the underlying levels of Ach, DA, and 5HIAA is excellent. My previous concerns have been adequately addressed.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

I appreciate the authors responses to my previous queries (and to the comments by other reviewers). The introduction does a better job contextualizing the data, and the additional details in the results and Methods sections help readers digest the material. I continue to think the topic is interesting and the manuscript is potentially impactful. However, I continue to be concerned about their analytical approaches and other aspects of the revised manuscript.

(a) Normalization

In my original review I wrote: "The normalization of neurochemical data seems unnecessary given the repeated-measures design of their analysis and could be problematic; by normalizing all data to the baseline data (p. 24), one artificially creates a baseline period with minimal variation (all are "0"; Figures 2, 3 & 5) that could inflate statistical power." I continue to feel that an analysis of normalized data that includes the baseline data is inappropriate because of the minimal variation in the normalized data for the baseline period. When the normalized data for the baseline period is included in the analysis, there is clearly variation in the extent of variability within each of the time periods (no variability at baseline, variability during periods 1 & 2; analogous to heteroscedasticity). For example, when analyzing the RAW DATA about the change in ACh release in experienced males listening to restraint vocalizations (thank you for releasing the raw data), there was a non-significant effect of time (baseline, period 1, and period 2; linear mixed effects model; F(2,12)=3.2, p=0.0793). However, when the normalized data for this dataset was analyzed (with baseline values being set at 100% for each mouse), there was a statistically significant effect (F(2,12)=4.5, p=0.0352). This example is just to illustrate how normalization can affect (e.g., inflate) statistical power.

That being said, I do think that it is reasonable to analyzed normalized data if the period used for normalization is NOT included in the analysis (see Figure 3 of one of the paper the authors listed in their response to reviewers: Galvez-Marquez et al., 2022). However, from the reading of this manuscript, it does seem like normalized baseline data are analyzed to assess how stimuli affect neurochemical concentrations.

We appreciate the reviewer’s point on the difference in variance caused by including the 100% baseline values in the analysis. After consulting with our statistician, we chose one of the two approaches suggested by the reviewer. Specifically, we reran the analysis to exclude the baseline and focus only on the playback windows and the group differences. The text in the results, the significance signs in the figures, and the discussion are corrected accordingly. Despite these changes, our major conclusions remains as before. We have included some descriptive statistics in the text because we think these are informative.

We decided to take this approach because the inter-individual variability in the raw data levels, caused by non-experimental factors, is too great to be useful. As we have stated before, these values are affected by probe placement, collection process, or differences in the HPLC or LC/MS runs. These effects are widely recognized in the field.

It is worth pointing out a few things about the papers listed by the authors. Li et al. (2023) does depict normalized microanalysis data but it isn't clear that any analysis of the normalized data is conducted. The same can be said about Holly et al. (2016). Further, in Bagley et al (2011), the authors depict normalized data in the figures but conduct analyses on the raw data ("After chronic morphine treatment, systemic naloxone injection increased GABA outflow in PAG by 41% (from 24.6 {plus minus} 2.9 nM to a peak of 34.8 {plus minus} 3.8 nM, n = 6, P = 0.016), but did not alter GABA levels after vehicle treatment (39.8 {plus minus} 8.3 to 38.6 {plus minus} 7.4 nM with naloxone at matched peak time, n = 4; Fig. 3a)". This latter approach (analyzing raw data in a repeated-measures manner and depicted normalized data) seems reasonable for the authors of the current study.

(b) Clarification and modification of statistical models

When analyzing the effect of experience on neuromodulator release, the authors analyze the experienced and inexperienced mice independently (e.g., figure 3 vs. 6). The ideal way to assess the effects of experience is to create a factorial model. For example, one could analyze a full factorial model with experience (exp vs. inexp), stimulus time (mating vs. restraint) and time (baseline, period 1 vs period 2, assuming raw data are used). If one wanted to exclude the baseline period because group differences in baseline are not informative, conducting a factorial analysis of normalized data with just the data from period 1 and 2 seems fine. I believe an analysis like this will help increase the legitimacy of the analysis. For example, when analyzing the normalized data (periods 1 and 2) of experienced and inexperienced males in response to mating or restraint vocalizations, you find a significant interaction between experience and stimulus type. Finding an effect of experience in an analysis that includes both experienced and inexperienced mice is ideal from an analytical framework.

In Figure 6, it is not clear what the statistical model is and what the interactions mean. For example, in the figure legend for figure 6, the authors report time*context and time*sex interactions. However, in this analysis there are two groups of inexperienced males (males that are listening to restraint vocalizations, males that are listening to mating vocalizations) and one group of females (females that are listening to mating vocalizations); in other words, this is an unbalanced analysis. So, when the authors indicate a time*context interaction, does that mean they are comparing the male-restraint group to the combination of males and females listening to mating vocalizations? And when they talk about a time*sex interaction, are they analyzing how males listening to either mating or restraint vocalizations differ from females listening to a mating vocalization? This all seems peculiar to me.

- A similar set of questions could be raised about interaction effects depicted in Figure 4.

Overall, I would like this manuscript to be reviewed by a statistician to provide additional input on how best to analyze the data.

We followed the reviewer’s suggestions to clarify the statistical model in studying the experience effect. After further consultation with the statistician, we reran the analysis on experience effect, including all the groups of EXP and INEXP animals together.

Design: Intercept + Sex +Context + Experience+ Sex* Experience + Context* Experience.

The model is not full factorial as recommended by the statistician, because we don’t have females in the restraint group and that would make an unbalanced design. Therefore, running GLM based on the above model and included factors, as advised by the statistician, is the best way of approaching the analysis for the current dataset.

We have corrected text in the figure captions, results, discussion, and data analysis sections of the manuscript related to the effect of experience and its interactions. The GLM models are clarified for all the figures in the “data analysis” section of the manuscript. We have clarified that the major effect of experience on neuromodulators was seen in the ACh data.

(c) Analysis of post-stimulus period

I agree with Reviewer 3 that analyzing the post-stimulus period would be useful. As mentioned in the original review, these data could serve as an opportunity to show that the neurochemical levels returned to baseline and add further support for the model described in Figure 6. In addition, these data could help reveal the link between neurochemical release, auditory responses, and behavior. If neurochemical changes reflect auditory responses, then these should back to baseline during the post-stimulus period. In addition, if behavioral variation (e.g., between mice hearing mating vs. restraint stimuli) persists following the termination of playback, then one could similarly assess whether neurochemical variation persists following playback. If the latter is the case, then the neurochemical release could be more related to the behavior than to the playback stimulus itself.

We did not change this analysis. Our response to Reviewer 3’s comment is shown below.

“We decided not to include analyses of the post-stimulus period because this period is subject to wider individual and neuromodulator-specific effects and because it weakens statistical power in addressing the core question—the change in neuromodulator release DURING vocal playback. We agree that the general question is of interest to the field, but we don’t think our study is best designed to answer that question.”

This was accepted by Reviewer 3. We also note that release patterns have multiple time courses (e.g., Aitta-aho et al., 2018 for ACh), and thus may not support an assumption that levels should return to baseline shortly after playback offset.

Minor comments:

Page 7, line 15: I suggest changing "vocalization-dependent" to "stimulus-dependent" because the former could connote patterns of release related to the animal itself vocalizing.

Changed to: “There were also distinct patterns of ACh and DA release into the BLA depending on the type of vocalization playback (Fig 3C,D).”

Discussion section: The authors should point out a few caveats with their experiments in the Discussion section. First, experienced animals received both mating (social) and restraint experiences, and it is not clear to what degree each type of experience affected neural and behavioral responses (i.e., specificity of experience effects). For example, mating experience can lead to a wide range of physiological changes, including a resilience to stress (e.g., Leuner et al., PLoS One, 2010; Arnold et al., Hormones and Behavior, 2019), so it is possible that mating experiences by themselves could have induced these changes. Or it could be that experiencing restraint stress affects responses to mating stimuli. This could be added to the first paragraph in page 16. (The authors could also discuss which aspects of the sexual encounters might be most important for the behavioral and neural plasticity.)

We have added text to raise this issue, stating that it is unknown wither the experience effects are specific and citing the above references concerning the generalized effects of certain experiences.

Discussion section: It would also be useful for the authors to discuss the extent to which behavior might be driving the neurochemical changes. Some of the analyses suggest that the release is independent of the behavior (e.g., reflects a sensory responses) but this could be emphasized more in the Discussion.

We believe that we have addressed this issue sufficiently in our previous response to related issues raised by this reviewer. As we note, there are limitations in the time resolution of microdialysis data that render the suggested discussion highly speculative. We plan to use other methods to assess this in future experiments.

Figure 2, legend: Please note that the text above the images describes the stimulus played back to these animals and their hormonal state, and not the type of experienced they underwent (i.e., clarify the titles)

Changed as requested.

I also agree with Reviewer 3 that "mating experience" is a misnomer for this manuscript. "Social experience with a female" is a more accurate descriptor. If they wanted to specifically provide mating experience, males should have only been tested with estrus (receptive females). I don't think this wording change detracts from their findings.

We have not changed this term. As noted in our previous response to Reviewer #3, we stated: “In the mating experience, mounting or attempted mounting was required for the animal to be included in subsequent testing.” Due to this requirement, the term “mating behavior” is informative and appropriate. In our view, “Social experience with a female” does not adequately describe our inclusion criterion or the experience.

Reviewer #3 (Recommendations For The Authors):

The work by Ghasemahmad et al. has the potential to significantly advance our understanding of how neuromodulators provide internal-state signals to the basolateral amygdala (BLA) while an animal listens to social vocalizations.

Ghasemahmad et al. made changes to the manuscript that have significantly improved the work. In particular, the transparency in showing the underlying levels of Ach, DA, and 5HIAA is excellent. My previous concerns have been adequately addressed. I only have a few minor suggestions for the text and one figure.

Minor suggestions:

Page 2, Ln 9: add adult before male and female mice

Changed as requested

Page 4, Ln 10: add a period after Tsukano et al., 2019)

Changed as requested

Page 6, Ln 9: what did you mean by "their interaction"? Being more specific, but concise, would help the readers.

We revised the wording to clarify that the neuromodulatory systems interact in the emission of positive and negative vocalizations.

Page 6, Ln 17: You mention Stim 1 and Stim 2, but the stimuli are not defined at this point. The clear explanation is provided in the following paragraph. Maybe consider switching the order and define the stimuli before you describe the liquid chromatography/mass spectrometry technique.

We have revised and merged these paragraphs so that Stim 1 and Stim 2 are defined on first use. We also revised our description of the depiction and analysis of neurochemical data.

Page 11, Ln 12: replace well-proven with well-documented

Changed as requested

Figure 2: There are two arrows pointing towards a single track. I assume one of the arrows is a duplicate. If so, delete one of the arrows. If not, please explain what the second arrow represents.

Arrow removed

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The basolateral amygdala (BLA), a brain center of emotional expression, contributes to acoustic communication by first interpreting the meaning of social sounds in the context of the listener’s internal state, then organizing the appropriate behavioral responses. We propose that modulatory neurochemicals such as acetylcholine (ACh) and dopamine (DA) provide internal-state signals to the BLA while an animal listens to social vocalizations. We tested this in a vocal playback experiment utilizing highly affective vocal sequences associated with either mating or restraint, then sampled and analyzed fluids within the BLA for a broad range of neurochemicals and observed behavioral responses of male and female mice. In male mice, playback of restraint vocalizations increased ACh release and usually decreased DA release, while playback of mating sequences evoked the opposite neurochemical release patterns. In non-estrus female mice, patterns of ACh and DA release with mating playback were similar to males. Estrus females, however, showed increased ACh, associated with vigilance, as well as increased DA, associated with reward-seeking. Experimental groups that showed increased ACh release also showed the largest increases in an aversive behavior. These neurochemical release patterns and several behavioral responses depended on a single prior experience with the mating and restraint behaviors. Our results support a model in which ACh and DA provide contextual information to sound analyzing BLA neurons that modulate their output to downstream brain regions controlling behavioral responses to social vocalizations.

Revised version: The basolateral amygdala (BLA), a brain center of emotional expression, contributes to acoustic communication by first interpreting the meaning of social sounds in the context of the listener’s internal state, then organizing the appropriate behavioral responses. We propose that modulatory neurochemicals such as acetylcholine (ACh) and dopamine (DA) provide internal-state signals to the BLA while an animal listens to social vocalizations. We tested this in a vocal playback experiment utilizing highly affective vocal sequences associated with either mating or restraint, then sampled and analyzed fluids within the BLA for a broad range of neurochemicals and observed behavioral responses of adult male and female mice. In male mice, playback of restraint vocalizations increased ACh release and usually decreased DA release, while playback of mating sequences evoked the opposite neurochemical release patterns. In non-estrus female mice, patterns of ACh and DA release with mating playback were similar to males. Estrus females, however, showed increased ACh, associated with vigilance, as well as increased DA, associated with reward-seeking. Experimental groups that showed increased ACh release also showed the largest increases in an aversive behavior. These neurochemical release patterns and several behavioral responses depended on a single prior experience with the mating and restraint behaviors. Our results support a model in which ACh and DA provide contextual information to sound analyzing BLA neurons that modulate their output to downstream brain regions controlling behavioral responses to social vocalizations.

Introduction

Added in the revised version: In social interactions utilizing vocal communication signals, the acoustic features of the signals carry emotional information (Altenmüller et al., 2013; Darwin, 1872; Seyfarth & Cheney, 2003). Listeners receive and analyze the acoustic information, compare it with previous experiences, identify the salience and valence of such information, and respond with appropriate behaviors.

Earlier version: In social interactions utilizing vocal communication signals, the acoustic features of the signals carry emotional information (Altenmüller et al., 2013; Darwin, 1872; Seyfarth & Cheney, 2003). Listeners receive and analyze the acoustic information, compare it with previous experiences, identify the salience and valence of such information, and respond with appropriate behaviors. These integrated functions depend on brain circuits that include the amygdala, a region located within the temporal lobe that is recognized to play a role in orchestrating emotional responses to salient sensory stimuli (J. E. LeDoux, 2000; McGaugh, 2002; Sah et al., 2003). The amygdalar target of auditory inputs from the thalamus and cortex is the basolateral amygdala (BLA) (J. LeDoux et al., 1984; Romanski & LeDoux, 1993; Shi & Cassell, 1997; Tsukano et al., 2019) By integrating this auditory input with other sensory inputs and inputs from other limbic areas, BLA neurons shape appropriate behavioral responses (Beyeler et al., 2016; Gründemann et al., 2019; Namburi et al., 2015, 2016) via projections to downstream targets such as the nucleus accumbens (Ambroggi et al., 2008; Stuber et al., 2011) and central nucleus of the amygdala (Ciocchi et al., 2010).

Revised version: These integrated functions depend on brain circuits that include the amygdala, a region located within the temporal lobe that is recognized to play a role in orchestrating emotional responses to salient sensory stimuli (J. E. LeDoux, 2000; McGaugh, 2002; Sah et al., 2003). The amygdalar target of auditory inputs from the thalamus and cortex is the basolateral amygdala (BLA) (J. LeDoux et al., 1984; Romanski & LeDoux, 1993; Shi & Cassell, 1997; Tsukano et al., 2019). By integrating this auditory input with other sensory inputs and inputs from other limbic areas, BLA neurons shape appropriate behavioral responses (Beyeler et al., 2016; Gründemann et al., 2019; Namburi et al., 2015, 2016) via projections to downstream targets such as the nucleus accumbens (Ambroggi et al., 2008; Stuber et al., 2011) and central nucleus of the amygdala (Ciocchi et al., 2010).

The BLA processes vocal and other acoustic information in a context-dependent manner (Gadziola et al., 2016; Grimsley et al., 2013; Matsumoto et al., 2016; Wenstrup et al., 2020; Wiethoff et al., 2009). Contextual information may arise from inputs associated with other sensory modalities (e.g., somatic sensation or olfaction) (Grimsley et al., 2013; Lanuza et al., 2004; McDonald, 1998), but in other cases the contextual information is associated with an animal’s internal state. Sources of internal state cues to BLA include brain circuits involving modulatory neurochemicals (i.e., neuromodulators), known to affect the processing of sensory signals, thus shaping attention, emotion, and goal-directed behaviors (Bargmann, 2012; Likhtik & Johansen, 2019; Schofield & Hurley, 2018). While previous work has demonstrated a role for some neuromodulators—dopamine (DA) and acetylcholine (ACh)—in the production of social vocalizations (Rojas-Carvajal et al., 2022; Silkstone & Brudzynski, 2020), it remains unclear how these and other neuromodulators contribute to vocal processing in social interactions.

This study investigates contextual information provided by neuromodulatory inputs to the BLA in response to vocal communication signals. Our hypothesis is that these salient vocalizations elicit distinct patterns of neuromodulator release into the BLA, by which they shape the processing of subsequent meaningful sensory information. We further hypothesize that these neuromodulatory patterns may depend on longer term processes that are critical to vocal communication: to experience with these behaviors and the accompanying vocalizations, to sex, and to estrous stage in females. To test these hypotheses, we conducted playback experiments in a mouse model to understand the behavioral and neuromodulator responses to salient vocalizations associated with very different behavioral states.

Results

To study how vocalizations affect behaviors and release of neurochemicals within BLA, we first developed highly salient vocal stimuli associated with appetitive (mating) and aversive (restraint) behaviors of CBA/CaJ mice. From more intense, mating-related interactions between adult male and female mice that included female head-sniffing and attempted or actual mounting (Gaub et al., 2016; Ghasemahmad, 2020, see Materials and Methods), we selected several sequences of vocalizations to form a 20-minute mating vocal stimulus. These sequences included ultrasonic vocalizations (USVs) with harmonics, steps, and complex structure, mostly emitted by males, and low frequency harmonic calls (LFHs) emitted by females (Fig. 1A,C) (Finton et al., 2017; Gaub et al., 2016; Ghasemahmad, 2020; Hanson & Hurley, 2012). During short periods of restraint, mice produce distinctive mid-frequency vocalizations (MFVs) that are associated with anxiety-related behaviors and increased release of the stress hormone corticosterone (Dornellas et al., 2021; Grimsley et al., 2016; Niemczura et al., 2020). From vocal sequences produced by restrained mice, we created a 20-minute vocal stimulus, primarily containing MFVs and fewer USV and LFH syllables (Fig. 1B,C).

Behavioral / microdialysis experiments test how playback of affective vocal signals alters behaviors and neuromodulator release into the BLA.

A. Short sample of mating vocal sequence used in playback experiments. Recording was obtained during high-intensity mating interactions and included ultrasonic vocalizations (USVs), likely emitted by the male, as well as low frequency harmonic (LFH) calls likely emitted by the female. B. Short sample of restraint vocal sequence used in playback experiments. Recording was obtained from an isolated, restrained mouse (see Material and Methods) and consisted mostly of mid-frequency vocalization (MFV) syllables. C. Syllable types in mating playback sequences differ substantially from those in restraint playback sequences. Percentages indicate frequency-of-occurrence of syllable types across all examplars used in mating or restraint vocal stimuli (nMating = 545, nRestraint = 622 syllables). D. Experimental design in playback experiment. Days 1 and 2: each animal experienced restraint and mating behaviors (counterbalanced order across subjects). Day 2: a microdialysis guide tube was implanted in the brain above the BLA. Day 6: the microdialysis probe was inserted through guide tube into the BLA. Playback experiments began after several hours of habituation/equilibration. Behavioral observations and microdialysis sampling were obtained before, during, and after playback of one vocalization type. Microdialysis samples were analyzed using LC/MS method described in Materials and Methods. E. Schematic illustration of detailed sequencing of vocal stimuli, shown here for mating playback. A 20-min period of vocal playback was formed by seven repeated stimulus blocks of 170 s. The stimulus blocks were composed of five vocal exemplars (each represented by a different color) of variable length, with each exemplar followed by an equal duration of silence. Stimuli during the Stim1 and Stim 2 playback windows thus included identical blocks but in slightly altered patterns. See Material and Methods for in-depth description of vocalization playback.

Earlier version: We next asked whether these salient vocal stimuli, associated with very different behavioral states, could elicit distinct behaviors and patterns of neuromodulator release into the BLA. We focused on the neuromodulators acetylcholine (ACh) and dopamine (DA), since previous work suggests their interaction in emission of positive and negative vocalizations (Rojas-Carvajal et al., 2022; Silkstone & Brudzynski, 2020). Our experiments combined playback of the vocal stimuli, behavioral tracking and observations (see Table 1 for descriptions), and microdialysis of BLA extracellular fluid in freely moving mice (Fig. 1D). Fluids were analyzed using a liquid chromatography/mass spectrometry (LC/MS) technique that allowed simultaneous measurement of several neurochemicals and their metabolites in the same dialysate samples, including ACh, DA and the serotonin metabolite 5-HIAA (see Materials and Methods). All neurochemical results during Stim 1 and Stim 2 periods are expressed as percent change from the Pre-Stim control period. However, raw values of ACh, DA, and 5-HIAA are reported in supplementary tables 1, 2, and 3, respectively.

Revised version: We next asked whether these salient vocal stimuli, associated with very different behavioral states, could elicit distinct behaviors and patterns of neuromodulator release into the BLA. We focused on the neuromodulators acetylcholine (ACh) and dopamine (DA), since previous work suggests that these neuromodulatory systems interact in the emission of positive and negative vocalizations (Rojas-Carvajal et al., 2022; Silkstone & Brudzynski, 2020). Our experiments combined playback of the vocal stimuli, behavioral tracking and observations (see Table 1 for descriptions), and microdialysis of BLA extracellular fluid in freely moving mice (Fig. 1D).

Earlier version: Prior to the study, male and female mouse subjects had no experience with sexual or restraint behaviors. On the first two days of the experiment, mice in an experienced group (EXP, n=31) were each exposed to 90-min sessions with mating and restraint behaviors in a counterbalanced design (Fig. 1D). Mice were then implanted with a guide cannula for microdialysis. On the playback/sample collection day (Day 6), a microdialysis probe was inserted into the guide cannula. After a 4-hour period of mouse habituation and probe equilibration, we recorded behavioral reactions and sampled extracellular fluid from the BLA before (Pre-Stim) and during a 20 min playback period, divided into two 10-min stimulation/collection/observation periods that are designated Stim 1 and Stim 2 (Fig. 1D,E. See also Materials and Methods). Each mouse received playback of either the mating or restraint stimuli, but not both: same-day presentation of both stimuli would require excessively long playback sessions, the condition of the same probe would likely change on subsequent days, and quality of a second implanted probe on a subsequent day was uncertain. Data are reported only from mice with more than 75% of the microdialysis probe implanted within the BLA (Fig. 2).

Revised version: Prior to the study, male and female mouse subjects had no experience with sexual or restraint behaviors. On the first two days of the experiment, mice in an experienced group (EXP, n=31) were each exposed to 90-min sessions with mating and restraint behaviors in a counterbalanced design (Fig. 1D). Mice were then implanted with a guide cannula for microdialysis. On the playback/sample collection day (Day 6), a microdialysis probe was inserted into the guide cannula. After a 4-hour period of mouse habituation and probe equilibration, we recorded behavioral reactions and sampled extracellular fluid from the BLA before (Pre-Stim) and during a 20 min playback period, divided into two 10-min stimulation/collection/observation periods that are designated Stim 1 and Stim 2 (Fig. 1D,E; see also Materials and Methods). Each mouse received playback of either the mating or restraint stimuli, but not both: same-day presentation of both stimuli would require excessively long playback sessions, the condition of the same probe would likely change on subsequent days, and quality of a second implanted probe on a subsequent day was uncertain. Fluids were analyzed using a liquid chromatography/mass spectrometry (LC/MS) technique that allowed simultaneous measurement of several neurochemicals and their metabolites in the same dialysate samples, including ACh, DA, and the serotonin metabolite 5-HIAA (see Materials and Methods). All neurochemical results during Stim 1 and Stim 2 periods are expressed as a percentage relative to the Pre-Stim control period. However, raw values of ACh, DA, and 5-HIAA are reported in supplementary tables 1, 2, and 3, respectively. Data are reported only from mice with more than 75% of the microdialysis probe implanted within the BLA (Fig. 2).

Microdialysis probe locations for EXP and INEXP Groups.

Earlier version: Colored lines indicate recovered probe tracks that resulted from infusion of fluorescent tracers. Black solid lines indicate external capsule; black dashed lines indicate major amygdalar subdivisions. Arrows indicate tracks related to insets. Insets: photomicrographs shows dextran-fluorescein labeling that marks the placement of the microdialysis probes for two cases, 1340 and 1404. Abbreviations: B, basal nucleus of amygdala; CeA, central nucleus of amygdala; La, lateral nucleus of amygdala.

Revised version: Labels above BLA outlines indicate groups based on playback type, sex, and hormonal state. Colored lines indicate recovered probe tracks that resulted from infusion of fluorescent tracers. Black solid lines indicate external capsule; black dashed lines indicate major amygdalar subdivisions. Arrows indicate tracks related to insets. Insets: photomicrographs shows dextran-fluorescein labeling that marks the placement of the microdialysis probes for two cases, 1340 and 1404. Abbreviations: B, basal nucleus of amygdala; CeA, central nucleus of amygdala; La, lateral nucleus of amygdala.

Earlier version: We first describe tests to examine whether playback of mating and restraint vocalizations results in different behavioral and neurochemical responses in male mice. We observed that two behaviors, still-and-alert and flinching, displayed pronounced increases with restraint playback that were not observed in the mating playback group (Fig. 3A,B). There were also distinct, vocalization-dependent patterns of ACh and DA release into the BLA (Fig 3C,D). Thus, in response to restraint vocalizations, we observed a significant increase in ACh concentration during both Stim 1 and Stim 2 playback windows. Mating vocalizations, however, resulted in a significant decrease in ACh release during both playback periods. (Fig. 3C). DA release displayed opposite patterns to ACh, increasing significantly during Stim 2 playback of mating vocalizations but showing a decreasing pattern (although nonsignificant) during playback of restraint vocal sequences (Fig. 3D). For ACh, levels differed significantly between the mating and restraint groups during both Stim 1 and Stim 2, while DA levels differed significantly only during Stim 2 (Fig. 3C,D). In contrast, the serotonin metabolite 5-HIAA showed no significant change over time following playback of either vocal stimulus, nor significant differences between groups (Fig. 3E). These findings suggest that both behavioral responses and ACh and DA release are modulated in listening male mice by the affective content of social vocalizations.

Revised version: We first describe tests to examine whether playback of mating and restraint vocalizations results in different behavioral and neurochemical responses in male mice. We observed that two behaviors, still-and-alert and flinching, showed differential effects of playback type, increasing during restraint playback relative to mating playback during the Stim 1 period (Fig. 3A, B). There were also distinct patterns of ACh and DA release into the BLA depending on the type of vocalization playback (Fig 3C, D). Thus, in response to restraint vocalizations, we observed an increase in ACh concentration during both Stim 1 (M ± SD: +24 ± 31% re Pre-Stim baseline) and Stim 2 (+25 ± 16%) playback windows. Mating vocalizations, however, resulted in a decrease in ACh release during both playback periods (Stim 1: -19 ± 21.0%; Stim 2: -16.4 ± 24.0%). (Fig. 3C). DA release displayed opposite patterns to ACh, increasing during playback of mating vocalizations during Stim 1 (+12.0 ± 27.4%) and Stim 2 (+25 ± 39% re baseline), but decreasing during playback of restraint vocal sequences (Stim 1: -11.0± 11.2%; Stim 2: -22 ± 25.6%) (Fig. 3D). For ACh, levels differed significantly between the mating and restraint groups during both Stim 1 and Stim 2, while DA levels differed significantly only during Stim 2 (Fig. 3C, D). In contrast, the serotonin metabolite 5-HIAA showed no distinct pattern over time following playback of either vocal stimulus, nor significant differences between groups (Fig. 3E). These findings suggest that both behavioral responses and ACh and DA release are modulated in listening male mice by the affective content of social vocalizations.

Behavioral and neuromodulator responses to vocal playback in male mice differ by behavioral context of vocalizations.

Earlier version: A, B. Boxplots show number of occurrences of specified behavior in 10-min observation periods before (Pre-Stim) and during (Stim 1, Stim 2) playback of mating or restraint vocal sequences (nMating = 7, nRestraint = 6). Note that playback sequences during Stim 1 and Stim 2 periods were identical within each group. During playback, the restraint group increased still-and-alert behavior compared to the mating group (A) (context: F (1,11) = 9.6, p=0.01, ŋ2=0.5), and increased flinching behavior (B) (time*context: F (1.1,12.2) =6.3, p=0.02, ŋ2=0.4) compared to the Pre-Stim baseline and to the mating group. C-E. Boxplots show differential release of acetylcholine (ACh), dopamine (DA), and 5-HIAA relative to the Pre-Stim period, during mating and restraint vocal playback (nMating = 9, nRestraint = 7). C. Significant differences in ACh for restraint (increase) playback in Stim 1 and Stim 2 vs Pre-Stim baseline and vs males in mating playback (time*context: F (2,28) = 6.6,p=0.004, ŋ2=0.32). D. Significant differences for DA for mating (increase) over time and vs restraint playback (time*context: F (2,28) = 6.9, p=0.004, ŋ2=0.33). E. No significant changes in 5-HIAA during vocal playback in male mice (time*context: F (2,28) = 0.97, p=0.4, ŋ2 = 0.06). A-E. Statistical testing examined all time windows within groups and all intergroup comparisons within time windows. Only significant tests are shown. Repeated measures GLM: *p<0.05, **p<0.01, ***p<0.001 (Bonferroni post hoc test). Time windows comparison: 95% confidence intervals. See Data Analysis section in Materials and Methods for description of box plots.

Revised version: A, B. Boxplots show number of occurrences of specified behavior in 10-min observation periods before (Pre-Stim) and during (Stim 1, Stim 2) playback of mating or restraint vocal sequences (nMating = 7, nRestraint = 6). Note that playback sequences during Stim 1 and Stim 2 periods were identical within each group. During playback, the restraint group increased still-and-alert behavior compared to the mating group (A) (context: F (1,11) = 9.6, p=0.01, ŋ2=0.5), and increased flinching behavior (B) (time*context: F (1.1,12.2) =6.3, p=0.02, ŋ2=0.4) compared to the Pre-Stim baseline and to the mating group. C-E. Boxplots show differential release of acetylcholine (ACh), dopamine (DA), and 5-HIAA relative to the Pre-Stim period, during mating and restraint vocal playback (nMating = 9, nRestraint = 7). C. Significant differences in ACh for restraint (increase) playback in Stim 1 and Stim 2 vs males in mating playback (Main effect of context: F (1,14) = 22.6,p<0.001, ŋ2=0.62). D. Significant differences for DA for mating (increase) vs restraint playback (Main effect of context: F (1,14) = 7.4, p=0.02, ŋ2=0.35). E. No significant changes in 5-HIAA during vocal playback in male mice (context: F (1,14) = 1.36, p=0.3, ŋ2 = 0.09). A-E. Statistical testing examined time windows within groups for behavioral observations and all intergroup comparisons within time windows for both behavior and neuromodulator data. Only normalized neuromodulators in Stim 1 and Stim 2 were used for statistical comparisons. Only significant tests are shown. Repeated measures GLM: *p<0.05, **p<0.01, ***p<0.001 (Bonferroni post hoc test). Time windows comparison: 95% confidence intervals. See Data Analysis section in Materials and Methods for description of box plots.

As male and female mice emit different vocalizations during courtship and mating (Finton et al., 2017; Grimsley et al., 2013; Neunuebel et al., 2015; Sales (née Sewell), 1972), we tested whether playback of vocal interactions associated with mating (Fig. 1A) results in different behavioral and neurochemical responses in listening male and female mice. Since our testing included both estrus and non-estrus females, we further examined the estrous effect on neurochemical release and behavioral reactions.

Earlier version: Playback of mating vocalizations resulted in some general and sex-based differences in behavioral responses. For instance, all groups displayed increased attending behavior (Fig. 4A). In females, regardless of estrous stage, still-and-alert behaviors increased and rearing decreased relative to male mice (Fig. 4B,C). This change in motor activity was further supported by video tracking results showing a significant decrease in distance traveled by female mice in response to mating vocal playback (significant sex effect (F (1,15) =9.3, p=0.008, ŋ2=0.4). Another behavior differed by estrous stage during mating playback: females in estrus displayed a strikingly higher number of flinching behaviors compared to males and non-estrus females (Fig. 4D). Our analysis of neuromodulator responses to mating vocalization playback revealed an estrous-dependent modulation of ACh levels during playback. ACh concentration in estrus females increased significantly during both Stim 1 and Stim 2 periods, whereas ACh showed decreases in both males and non-estrus females that were significant in the Stim 2 period (Fig. 4E). Moreover, during playback, the ACh in estrus females was significantly higher than both males and non-estrus females. DA release showed an increase during mating playback across all three experimental groups, although the time course differed (Fig. 4F). Similar to male groups in restraint and mating playback, the 5-HIAA release patterns in females showed no modulation during mating vocal playback (Fig. 4G).

Revised version: Playback of mating vocalizations resulted in some general and some sex-based differences in behavioral responses. For instance, all groups displayed increased attending behavior (Fig. 4A). In females, regardless of estrous stage, rearing decreased and still-and-alert behavior increased significantly relative to male mice (Fig. 4B, C). This change in motor activity was further supported by video tracking results showing a significant decrease in distance traveled by female mice in response to mating vocal playback (significant sex effect (F (1,15) =9.3, p=0.008, ŋ2=0.4). However, one behavioral change was estrous dependent: females in estrus displayed a strikingly higher number of flinching behaviors compared to males and non-estrus females during both Stim 1 and Stim 2 periods (Fig. 4D). Our analysis of neuromodulator responses to mating vocalization playback revealed an estrous-dependent modulation of ACh levels during playback. ACh concentration in estrus females increased during both Stim 1 and Stim 2 periods, whereas ACh decreased in both males and non-estrus females (Fig. 4E). Moreover, during playback, post hoc comparison with Bonferroni correction showed that the ACh in estrus females was significantly higher (Stim 1, +29.3 ± 27.0 %; Stim 2, +25.5 ± 22%) than both males (Stim 1, -18 ± 21.0%; Stim 2, -16.0 ± 24.0%) and non-estrus females (Stim 1, -6.0 ± 14 %; Stim 2, -35.0 ± 29.6%). DA release showed increases during mating playback within all three experimental groups (Males: Stim 1, +12.0 ± 27.0%, Stim 2, +26.0 ± 39.1%;Estrus females: Stim 1, +24.1 ± 37.0%, Stim 2, +23.0 ± 27.0%; Non-estrus females: Stim 1, +27.0 ± 17.0%, Stim 2, +49.0 ± 48.0%), but no significant differences among groups (Fig. 4F). Similar to male groups in restraint and mating playback, the 5-HIAA release patterns in females showed no clear modulation pattern during mating vocal playback or differences among groups (Fig. 4G).

Behavioral and neuromodulator responses to playback of mating vocal sequences differ by sex and female estrous stage.

Earlier version: A-D. Occurrences of specified behaviors in 10-min periods before (Pre-Stim) and during (Stim 1, Stim 2) mating vocal playback (nMale = 7, nEstrous Fem = 6, nNon-estrous Fem = 5). A. Attending behavior increased during Stim 1 in response to mating vocal playback, regardless of sex or estrous stage (time*sex: F (2,30) = 0.12; p=0.9, ŋ2=0.008; time*estrous: F (2,30) = 1.1; p=0.4, ŋ2=0.07). B-C. Females regardless of estrous stage reared less (sex: F (1,15) =10.22, p=0.006; ŋ2=0.4; estrous: F (1,15) =0.2, p=0.7, ŋ2=0.01) and displayed more Still-and-Alert behaviors (Sex: F (1,15) =5.17, p=0.04, partial ŋ2= 0.3, estrous: (1,15) =0.07, p=0.8, partial ŋ2= 0.005) than males during mating vocal playback D. Estrus females, but not non-estrus females or males, showed a significant increase in flinching behavior during Stim 1 and Stim 2 periods (time*estrous: F (2,30) =9.0, p<0.001, ŋ2=0.4). E-G. Changes in concentration of acetylcholine (ACh), dopamine (DA), and 5-HIAA relative to the Pre-Stim period, evoked during Stim 1 and Stim 2 periods of vocal playback (nMale = 9, nEstrus Fem = 8, nNon-estrus Fem = 7). E. Release of ACh during mating playback increased in estrus females (Stim1, Stim 2) but decreased in males and non-estrus females (Stim 2). Among groups, there was a significant estrous effect (time*estrous: F (2,42) =9.9, p<0.001, ŋ2=0.32), increasing above Pre-Stim period in estrus females. F. DA release during mating playback increased in all groups relative to Pre-Stim period (time: F (2,42) = 10.9, p<0.001, ŋ2=0. 34), with no significant sex (F (1,21) =0.9, p=0.4, ŋ2=0.04) or estrous effect (F (1,21) =0.9, p=0.4, ŋ2=0.04). G. No significant changes in 5-HIAA during mating sequence playback (time*sex: F (2,42) = 0.08, p=0.9, ŋ2=0.004; time*estrus: F (2,42) =1.2, p=0.3, ŋ2 =0.05). A-G: Repeated measures GLM: *p<0.05, **p<0.01, ***p<0.001 (Bonferroni post hoc test). Time windows comparison: 95% confidence intervals.

Revised version: A-D. Occurrences of specified behaviors in 10-min periods before (Pre-Stim) and during (Stim 1, Stim 2) mating vocal playback (nMale = 7, nEstrus Fem = 6, nNon-estrus Fem = 5). A. Attending behavior increased during Stim 1 in response to mating vocal playback, regardless of sex or estrous stage (time: F (2,30) = 32.6, p<0.001, ŋ2=0.7; time*sex: F (2,30) = 0.12; p=0.9, ŋ2=0.008; time*estrous: F (2,30) = 1.1; p=0.4, ŋ2=0.07). B-C. Females regardless of estrous stage reared less (sex: F (1,15) =10.22, p=0.006; ŋ2=0.4; estrous: F (1,15) =0.2, p=0.7, ŋ2=0.01) and-displayed more Still-and-Alert behaviors (Sex: F (1,15) =5.2, p=0.04, partial ŋ2= 0.3, estrous: (1,15) =0.07, p=0.8, partial ŋ2= 0.005) than males during mating vocal playback. D. Estrus females, but not non-estrus females or males, showed a significant increase in flinching behavior during Stim 1 and Stim 2 periods (time*estrous: F (2,30) =9.0, p<0.001, ŋ2=0.4). E-G. Changes in concentration of acetylcholine (ACh), dopamine (DA), and 5-HIAA relative to the Pre-Stim period, evoked during Stim 1 and Stim 2 periods of vocal playback (nMale = 9, nEstrus Fem = 8, nNon-estrus Fem = 7). E. Release of ACh during mating playback increased in estrus females (Stim1, Stim 2) but decreased in males and non-estrus females (Stim 2). Among groups, there was a significant estrous effect (estrous: F (1,21) =29.0, p<0.001, ŋ2=0.0.6. F. DA release during mating playback increased in all groups relative to Pre-Stim period, with no significant sex (F (1,21) =0.8, p=0.4, ŋ2=0.04) or estrous effect (F (1,21) =0.9, p=0.4, ŋ2=0.04). G. No significant changes in 5-HIAA during mating sequence playback (sex: F (1,21) = 0.07, p=0.8, ŋ2 =0.004; estrus: F (1,21) =1.6, p=0.22, ŋ2 =0.07). A-G: Repeated measures GLM: *p<0.05, **p<0.01, ***p<0.001 (Bonferroni post hoc test). Time windows comparison: 95% confidence intervals.

Like male mice exposed to restraint vocalizations, estrus females showed robust and significant increases in flinching behavior in response to mating vocalizations. Both groups displayed increased ACh release. This supports the possible involvement of ACh in shaping such behavior in both males listening to restraint calls and in estrus females listening to mating vocalizations. Note that neuromodulator release, including ACh, has been previously linked to motor behaviors (Wall & Woolley, 2020), but the changes in ACh that we observed showed no relationship with behaviors involving motor activity such as rearing (restraint males: Stim 1: n = 6, r = 0.1, p = 0.8; Est females: n = 6, r = 0.7, p=0.08), locomotion (Stim 1: restraint males, n = 6, r = -0.2, p = 0.7; Est females: n = 6, r = 0.3, p = 0.6), or distance traveled (Stim 1: restraint males, n = 6, r = 0.16, p = 0.8; Est females: n = 6, r = -0.6, p = 0.2). This suggests that the observed changes in ACh reflect the valence of these vocalizations.

Earlier version: All EXP mice used in the above experiments had undergone a single session each to experience mating and restraint conditions prior to the playback session on Day 6 (Fig. 1D). Does such experience shape the release patterns of these neuromodulators in response to vocal playback? We tested male and female mice under identical vocal playback conditions as previous groups, except that they did not receive the restraint and mating experiences (INEXP groups). Since only one INEXP female was in a non-estrus stage during the playback session, our analysis of the effect of experience included only estrus females and males.

Revised version: All EXP mice used in the above experiments had undergone a single session each to experience mating and restraint conditions prior to the playback session on Day 6 (Fig. 1D). Does such experience shape the release patterns of these neuromodulators in response to vocal playback?

Added in the revised version: We tested male and female mice under identical vocal playback conditions as previous groups, except that they did not receive the restraint and mating experiences (INEXP groups). Since only one INEXP female was in a non-estrus stage during the playback session, our analysis of the effect of experience included only estrus females and males.

Earlier version: Several behavioral responses to vocalization playback differed between EXP and INEXP mice in a sex- or context-dependent manner. For example, EXP estrus females showed increased sound-evoked flinching behaviors compared to INEXP estrus females (Fig. 5A) in response to mating vocal sequences. These experience effects were not observed in males in response to mating or restraint vocal playback. Further, the differences in flinching behavior among EXP groups (male-mating vs male-restraint or estrus female-mating) were not evident among INEXP groups. For rearing behavior, EXP males responding to mating vocalizations showed an increase compared to INEXP males (Fig. 5B). This pattern was not observed during restraint playback for males or mating playback in estrus females. These findings indicate that behavioral responses to salient vocalizations result from interactions between sex of the listener or context of vocal stimuli with the previous behavioral experience associated with these vocalizations.

Revised version: Several behavioral responses to vocalization playback differed between EXP and INEXP mice in a sex- or context-dependent manner. For example, EXP estrus females showed significantly increased sound-evoked flinching behaviors compared to INEXP estrus females (Fig. 5A) in response to mating vocal sequences. These experience effects were not observed in males in response to mating or restraint vocal playback. Further, the differences in flinching behavior among EXP groups (male-mating vs male-restraint or estrus female-mating) were not evident among INEXP groups. For rearing behavior, EXP males responding to mating vocalizations showed an increase compared to INEXP males (Fig. 5B). This pattern was not observed during restraint playback for males or mating playback in estrus females. These findings indicate that behavioral responses to salient vocalizations result from interactions between sex of the listener or context of vocal stimuli with the previous behavioral experience associated with these vocalizations.

Earlier version: A single bout each of mating and restraint experience altered behaviors in response to vocal playback.

Revised version: A single bout each of mating and restraint experience altered behaviors and ACh release in response to vocal playback.

Earlier version: In all graphs, dots represent measures from individual animals obtained during the Stim 1 playback period; connected thick lines represent mean values across subjects. There were no differences in Pre-Stim values between INEXP and EXP mice for any group (e.g., males-restraint, males-mating, estrus females-mating. A. Experience increased flinching responses in estrus female mice but not in males in mating or restraint vocal playback (Females: nEXP = 6, nINEXP = 8, t=5.0, p<0.001; Male-mating: nEXP = 7, nINEXP = 7, t= 0.4, p=0.7; Male-restraint: nEXP = 6, nINEXP = 7, t=0.6, p=0.6); INEXP male mating vs restraint: t=2.08, p=0.06; mating female vs male: t=0.8, p=0.4; EXP male mating vs restraint: t= 2.5, p=0.03; mating female vs male: t=4.3, p<0.001. B. Experience increased rearing responses in males exposed to mating playback (Male-mating: t=3.2, p=0.007), but not in females (t=1.4, p=0.2) or males in restraint playback (t=0.4, p=0.7); INEXP male mating vs restraint: t=1.07, p=0.3, mating female vs male: t=1.7, p=0.1; EXP male mating vs restraint: t=2.0, p=0.08; mating female vs male: t=3.4, p=0.004. GLM repeated measures with Bonferroni post hoc test: *p<0.05, **p<0.01, ***p<0.001.

Revised version: In all graphs, dots represent measures from individual animals obtained during the Stim 1 playback period; thick horizontal lines represent mean values across subjects. There were no differences in behavioral counts for Pre-Stim values between INEXP and EXP mice for any group (e.g., males-restraint, males-mating, estrus females-mating. Neuromodulator values are normalized to the baseline level. A. Experience increased flinching responses in estrus female mice but not in males in mating or restraint vocal playback (time*sex*experience: F(1.6, 56) = 4.1, p=0.03, ŋ2 = 0.11)). B. Experience increased rearing responses in males exposed to mating playback (sex*experience: F(1,35)=5.3, P=0.03, ŋ2 =0.13). C. Estrus female mice -(sex*experience: F(1, 39) = 8.0, p=0.008, ŋ2 = 0.2,) and restraint male mice (context*experience = F(1, 39) = 13.0, p<0.001, ŋ2 = 0.2)displayed consistent experience effect for changes in ACh. D. DA did not show the EXP effect observed in ACh during vocal playback (sex*experience: F(1, 39) =0.12, p=0.7, ŋ2 = 0.003; context*experience = F(1, 39) = 4.0, p=0.052, ŋ2 =0.09. GLM repeated measures with Bonferroni post hoc test: *p<0.05, **p<0.01, ***p<0.001.

Added in the revised version: A major finding is that prior experience with mating and restraint behaviors shaped patterns of ACh release in response to vocal playback; there was a significant interaction of context and experience for ACh (F(1, 39) = 12.7, P<0.001; Figs. 5C, 6A). Thus, ACh release was significantly different between INEXP and EXP males in restraint playback groups. A similar pattern was observed in INEXP estrus females in response to mating vocal playback, i.e., the concentration of ACh re baseline failed to show such pronounced increases as observed in EXP females and showed a significant experience effect (Sex*experience interaction: F (1, 39) = 7.8, P= 0.008; Figs. 5C, 6A). Further, INEXP males in mating playback group failed to show the reduced ACh concentration observed in EXP males (INEXP Stim 1: +9.6 ± 15.0%; INEXP Stim 2: +6.1 ± 30.0; EXP Stim 1: -19.0 ± 21.0%; EXP Stim 2: -16.0 ± 24.0 %), even though these differences were not significant (Figs. 5C, 6A).

Removed in the revised version: A major finding is that prior experience with mating and restraint behaviors shaped patterns of ACh and DA release in response to vocal playback (Fig. 6). Thus, male and female mice lacking in previous mating and restraint experiences showed no significant change in ACh or DA concentrations in response to either vocal playback type (Fig. 6A,B). Further, there were no significant differences across groups during either Stim 1 or Stim 2 periods. These results contrast sharply with those from all EXP groups, in which release of ACh, DA, or both changed significantly during playback (Figs. 3C, 3D, 4E, 4F). Moreover, 5-HIAA concentrations, which were unaffected by sex, estrous stage, or playback type (Figs. 3E, 4G), were also unaffected by experience (Fig. 6C). Collectively, these data suggest that the playback vocalization type and estrous effects observed in neuromodulator release patterns and behavioral reactions depend on previous experience with the corresponding behaviors.

Earlier version: In INEXP mice, vocal playback failed to evoke significant changes in neuromodulator release.

Revised version: In INEXP mice, vocal playback failed to evoke distinct patterns in neuromodulator release.

Earlier version: Boxplots show change in concentration of the specified neuromodulator in 10-min playback periods (Stim 1, Stim 2) compared to the baseline level (Pre-Stim). Male-restraint, nINEXP = 7; Male-mating, nINEXP = 7; Estrus Female-mating, nINEXP = 7). A. ACh release shown for both playback periods (mating vs restraint: time*context: F (2,24) =0.56; p=0.6; male vs female: time*sex: F (2,24) =0.06; p=0.9). B. DA release shown for both playback periods (mating vs restraint: time*context: F (2,24) = 0.2; p=0.8; male vs female: time*sex: F (2,24) = 0.03; p=0.97). C. 5-HIAA concentration shown for both playback periods (mating vs restraint: time: F (1.3,15.6) =0.005, p=0.97; time*context: F (1.3,15.6) = 2.1; p=0.17; male vs female: time: F (1.3,16.2) =0.8; p=0.4; time* sex: F (1.3,16.2) = 0.5; p=0.55). A, B. Statistical testing examined all time windows within groups and all intergroup comparisons within time windows. No tests were significant. Repeated measures GLM (Bonferroni post hoc test). Time windows comparison: 95% confidence intervals.

Revised version: Boxplots show change in concentration of the specified neuromodulator in 10-min playback periods (Stim 1, Stim 2) compared to the baseline level (Pre-Stim). Male-restraint, nINEXP = 7; Male-mating, nINEXP = 7; Estrus Female-mating, nINEXP = 7). A. ACh release shown for both playback periods. B. DA release shown for both playback periods. C. 5-HIAA concentration shown for both playback periods.

Added in the revised version: DA release patterns, which were weaker in EXP groups compared to ACh patterns (Figs. 3C vs 3D; 4E vs 4F) showed no significant experience-dependent changes in any group (Fig. 5D). There was no evidence that vocal playback altered DA release in any of the three groups of INEXP mice (Fig 6B), unlike what was observed in EXP groups (Figs. 3D, 4F). 5-HIAA concentrations, which were unaffected by sex, estrous stage, or playback type (Figs. 3E, 4G), were also unaffected by experience (Fig. 6C).

Added in the revised version: Collectively, these data suggest that the playback vocalization type and estrous effects observed in ACh release patterns and behavioral reactions depend on previous experience with the corresponding behaviors.

Discussion

Earlier version: Functional imaging studies in humans and mechanistic studies in other species provide substantial evidence that the amygdala participates in circuits that process vocalizations (Frühholz et al., 2016; Liebenthal et al., 2016; Sander et al., 2003; Wenstrup et al., 2020; Voytenko et al., 2023), assess the valence of appetitive and aversive cues (Kyriazi et al., 2018; O’Neill et al., 2018; Pignatelli & Beyeler, 2019; Smith & Torregrossa, 2021), and shape appropriate behavioral responses to these cues (Gründemann et al., 2019; Lim et al., 2009; Schönfeld et al., 2020; Zhang & Li, 2018). Nonetheless, how contextual information is delivered to the amygdala and contributes to vocal processing is not well understood. Since the amygdala receives strong projections from neuromodulatory brain centers (Aitta-aho et al., 2018; Asan, 1997, 1998; Bigos et al., 2008; Bocchio et al., 2016; Carlsen et al., 1985), and since the role of these neurochemicals in providing internal state and contextual information is well-proven (Bocchio et al., 2016; Jiang et al., 2016; Likhtik & Johansen, 2019), we hypothesized that release patterns of neuromodulators into the BLA provide contextual information during processing of affective vocalizations. Our results show that these emotionally charged vocalizations result in distinct release patterns of acetylcholine (ACh) and dopamine (DA) into the BLA of male and female mice. Further, female hormonal state appears to influence ACh but not DA release into the BLA when processing mating vocalizations. Such context- or state-dependent changes were not observed in patterns of other neurochemicals (e.g., 5-HIAA). Finally, we showed that a single 90-min experience with intense behaviors is sufficient to establish strong, consistent patterns of neuromodulatory release into the amygdala. These data indicate that during analysis of affective vocalizations in the BLA, ACh and DA provide experience-dependent, state- and context-related information that can potentially modulate sensory processing within the BLA and thus shape an individual’s response to these vocalizations.

Revised version: Functional imaging studies in humans and mechanistic studies in other species provide substantial evidence that the amygdala participates in circuits that process vocalizations (Frühholz et al., 2016; Liebenthal et al., 2016; Sander et al., 2003; Wenstrup et al., 2020; Voytenko et al., 2023), assess the valence of appetitive and aversive cues (Kyriazi et al., 2018; O’Neill et al., 2018; Pignatelli & Beyeler, 2019; Smith & Torregrossa, 2021), and shape appropriate behavioral responses to these cues (Gründemann et al., 2019; Lim et al., 2009; Schönfeld et al., 2020; Zhang & Li, 2018). Nonetheless, how contextual information is delivered to the amygdala and contributes to vocal processing is not well understood. Since the amygdala receives strong projections from neuromodulatory brain centers (Aitta-aho et al., 2018; Asan, 1997, 1998; Bigos et al., 2008; Bocchio et al., 2016; Carlsen et al., 1985), and since the role of these neurochemicals in providing internal state and contextual information is well documented (Bocchio et al., 2016; Jiang et al., 2016; Likhtik & Johansen, 2019), we hypothesized that release patterns of neuromodulators into the BLA provide contextual information during processing of affective vocalizations. Our results show that these emotionally charged vocalizations result in distinct release patterns of acetylcholine (ACh) and dopamine (DA) into the BLA of male and female mice. Further, female hormonal state appears to influence ACh but not DA release into the BLA when processing mating vocalizations. Such context- or state-dependent changes were not observed in patterns of other neurochemicals (e.g., 5-HIAA), nor in the absence of experience with behaviors associated with the vocalizations. In particular, we showed that a single 90-min experience with intense behaviors is sufficient to establish strong, consistent patterns of ACh release into the amygdala. These data indicate that during analysis of affective vocalizations in the BLA, ACh and DA provide experience-dependent, state- and context-related information that can potentially modulate sensory processing within the BLA and thus shape an individual’s response to these vocalizations.

The linkage between neuromodulator release and behavioral responses to vocalizations varies. The strongest case is for a link between increased ACh release and flinching behavior. Males in the restraint playback group and estrus females in the mating playback group both displayed significantly increased ACh release and flinching behavior, occurring in both playback periods. This suggests a mechanistic relationship. Other significant behavioral responses to vocal stimuli did not closely match the timing of neurochemical changes, suggesting a weak mechanistic link. Overall, it is not surprising that most behaviors and neurochemicals do not match well, given the poor temporal resolution of the neurochemical sampling. We believe that our results show the need for a finer analysis of DA and ACh release based on techniques for high temporal resolution of neurochemical measurements (e.g., genetically encoded ACh indicator (Jing et al., 2018)), coupled with short interval behavioral observations. Only then can the role of the BLA and neuromodulator inputs in behavioral vocal responses be established more strongly. Further, we note that no conclusion can be made regarding relationships between noradrenalin and serotonin concentrations and vocal playback, since these were not detected by the neurochemical analysis.

Earlier version: The BLA receives strong cholinergic projections from the basal forebrain (Aitta-aho et al., 2018; Carlsen et al., 1985) that contribute to ACh-dependent processing of aversive cues and fear learning in the amygdala (Baysinger et al., 2012; Gorka et al., 2015; Jiang et al., 2016; Tingley et al., 2014). Our findings support these studies by demonstrating increased ACh release in BLA in response to playback of aversive vocalizations. Although the exact mechanism by which ACh affects vocal information processing in BLA is not clear yet, the result of ACh release onto BLA neurons seems to enhance arousal during emotional processing (Likhtik & Johansen, 2019). Our results, in conjunction with previous work, suggest mechanisms by which vocalizations affect ACh release and in turn drive behavioral responses (Fig. 7A). Cholinergic modulation in the BLA is mediated via muscarinic and nicotinic ACh receptors on BLA pyramidal neurons and inhibitory interneurons (Aitta-aho et al., 2018; Mesulam et al., 1983; Pidoplichko et al., 2013; Unal et al., 2015). During the processing of sensory information in the BLA, partially non-overlapping populations of neurons respond to cues related to positive or negative experiences (Namburi et al., 2015; Paton et al., 2006; Smith & Torregrossa, 2021). These neurons then project to different target areas involved in appetitive or aversive behaviors—the nucleus accumbens or central nucleus of the amygdala, respectively (Namburi et al., 2015).

Revised version: The BLA receives strong cholinergic projections from the basal forebrain (Aitta-aho et al., 2018; Carlsen et al., 1985) that contribute to ACh-dependent processing of aversive cues and fear learning in the amygdala (Baysinger et al., 2012; Gorka et al., 2015; Jiang et al., 2016; Tingley et al., 2014). Our findings support these studies by demonstrating increased ACh release in BLA in EXP animals in response to playback of aversive vocalizations. Although the exact mechanism by which ACh affects vocal information processing in BLA is not clear yet, the result of Ach release onto BLA neurons seems to enhance arousal during emotional processing (Likhtik & Johansen, 2019). Our results, in conjunction with previous work, suggest mechanisms by which vocalizations affect ACh release and in turn drive behavioral responses (Fig. 7A). Cholinergic modulation in the BLA is mediated via muscarinic and nicotinic ACh receptors on BLA pyramidal neurons and inhibitory interneurons (Aitta-aho et al., 2018; Mesulam et al., 1983; Pidoplichko et al., 2013; Unal et al., 2015). During the processing of sensory information in the BLA, partially non-overlapping populations of neurons respond to cues related to positive or negative experiences (Namburi et al., 2015; Paton et al., 2006; Smith & Torregrossa, 2021). These neurons then project to different target areas involved in appetitive or aversive behaviors—the nucleus accumbens or central nucleus of the amygdala, respectively (Namburi et al., 2015).

Proposed model for neuromodulation of salient vocalization processing via acetylcholine (ACh) and dopamine (DA) in the basolateral amygdala (BLA).

A. Cholinergic modulation of CeA-projecting neurons during aversive vocalization cue processing in the BLA. In the presence of aversive cues, ACh released from the basal forebrain acts on M1 ACh receptors (M1 mAChRs) to enhance the cue-induced excitatory responses of CeA -projecting neurons. In contrast, NAc-projecting neurons are quiescent, because they do not respond to aversive cues and are inhibited by interneurons that are activated through nicotinic ACh receptors (nAChRs). B. Dopaminergic modulation and enhancement of signal-to-noise ratio in response to reward-associated cues (appetitive vocalizations). When vocalizations or other rewarding cues are present, release of DA from VTA enhances D2R-mediated excitation in NAc-projecting neurons that are responsive to positive cues. In contrast, CeA -projecting neurons are not responsive to rewarding vocalizations and are inhibited by local interneurons. DA is thought to act on D1 DA receptors (D1Rs) in these local interneurons to shape a direct inhibition onto CeA-projecting neurons. Abbreviations: ACx: auditory cortex; BF, basal forebrain; CeA, central nucleus of amygdala; MGB, medial geniculate body; NAc, nucleus accumbens; VTA, ventral tegmental area.

In response to an aversive cue or experience (Fig. 7A), released ACh affects BLA neurons according to their activity. If projection neurons are at rest, ACh may exert an inhibitory effect by activating nicotinic ACh receptors on local GABAergic interneurons, which in turn synapse onto the quiescent pyramidal neurons. This results in GABA-A mediated inhibitory postsynaptic potentials (IPSPs) in the pyramidal neurons. Alternately, direct activation of M1 ACh receptors on pyramidal neurons, activating inward rectifying K+ currents, may result in additional ACh-mediated inhibition (Fig. 7A) (Aitta-aho et al., 2018; Pidoplichko et al., 2013; Unal et al., 2015). When BLA pyramidal neurons are already active due to strong excitatory input associated with aversive cues, M1 receptor activation can result in long afterdepolarizations that produce persistent firing lasting as long as ACh is present (Jiang et al., 2016; Unal et al., 2015). Such a process may explain persistent firing observed in single neuron responses to aversive social vocalizations in bats (Gadziola et al., 2012; Peterson & Wenstrup, 2012). Through this process of inhibiting quiescent neurons and enhancing activation and persistent firing in active neurons, ACh sharpens the population signal-to-noise ratio (SNR) during the processing of salient, aversive signals in the BLA. These neurons, processing negative cues, likely project to the central nucleus of the amygdala (CeA) to regulate defensive behaviors such as escape and avoidance (Fig. 7A) (Beyeler et al., 2016; Davis et al., 2010; Namburi et al., 2015). In agreement, our behavioral findings show increased behaviors such as flinching along with increased release of ACh during processing of aversive vocalizations. Such prolonged afterdepolarizations provide the appropriate condition for associative synaptic plasticity (Likhtik & Johansen, 2019) that underlies an increase in the AMPA/NMDA current ratio in CeA-projecting neurons during processing of aversive cues.

Removed in the revised version: Dopaminergic innervation from the ventral tegmental area (VTA) (Asan, 1998) acts on BLA neurons via D1 and D2 receptors, both G-protein coupled receptors. Dopamine is important in reward processing, fear extinction, decision making, and motor control (Ambroggi et al., 2008; Di Ciano & Everitt, 2004; Lutas et al., 2019). We observed increased DA release in the BLA in response to mating vocalizations both for males and for females across estrous stages.

Earlier version: Electrophysiological studies show that DA enhances sensory processing in BLA neurons by increasing the population response SNR in a process like ACh (Kröner et al., 2005; Vander Weele et al., 2018). Thus, during processing of mating vocalizations or those related to other rewarding experiences, DA presence in the vicinity of BLA pyramidal neurons and interneurons is enhanced (Fig. 7B). For neurons with elevated spiking activity during processing of appetitive vocalizations or other sensory stimuli, DA acts on D2 receptors of pyramidal cells to further enhance neuronal firing and result in persistent firing of these projection neurons. Conversely, in BLA projection neurons that do not respond to such positive cues, such as CeA-projecting neurons, DA exerts a suppressive effect directly via D1 receptors and indirectly by activating inhibitory interneuron feedforward inhibition (Kröner et al., 2005). The net result of this process in response to appetitive vocalizations is an enhancement of activity in the reward-responding neurons and suppression of activity in aversive-responding neurons. This process likely depends on the increase in synaptic plasticity via an enhanced AMPA/NMDA current ratio during processing such cues in NAc-projecting neurons in the BLA (Namburi et al., 2015; Otani et al., 2003; van Vugt et al., 2020). Our findings suggest that this may occur in the BLA in response to appetitive vocalizations.

Revised version: Dopaminergic innervation from the ventral tegmental area (VTA) (Asan, 1998) acts on BLA neurons via D1 and D2 receptors, both G-protein coupled receptors. Dopamine is important in reward processing, fear extinction, decision making, and motor control (Ambroggi et al., 2008; Di Ciano & Everitt, 2004; Lutas et al., 2019). We observed increased DA release in the BLA in EXP groups in response to mating vocalizations both for males and for females across estrous stages. Electrophysiological studies show that DA enhances sensory processing in BLA neurons by increasing the population response SNR in a process like ACh (Kröner et al., 2005; Vander Weele et al., 2018). Thus, during processing of mating vocalizations or those related to other rewarding experiences, DA presence in the vicinity of BLA pyramidal neurons and interneurons is enhanced (Fig. 7B). For neurons with elevated spiking activity during processing of appetitive vocalizations or other sensory stimuli, DA acts on D2 receptors of pyramidal cells to further enhance neuronal firing and result in persistent firing of these projection neurons. Conversely, in BLA projection neurons that do not respond to such positive cues, such as CeA-projecting neurons, DA exerts a suppressive effect directly via D1 receptors and indirectly by activating inhibitory interneuron feedforward inhibition (Kröner et al., 2005). The net result of this process in response to appetitive vocalizations is an enhancement of activity in the reward-responding neurons and suppression of activity in aversive-responding neurons. This process likely depends on the increase in synaptic plasticity via an enhanced AMPA/NMDA current ratio during processing such cues in NAc-projecting neurons in the BLA (Namburi et al., 2015; Otani et al., 2003; van Vugt et al., 2020). Our findings suggest that this may occur in the BLA in response to appetitive vocalizations.

Earlier version: As the results with males listening to restraint vocalizations demonstrate, increased ACh release in BLA is associated with processing aversive cues. How, then, should the increased ACh release in estrus females during mating vocal playback be interpreted? Previous work shows that neuromodulation of amygdalar and other forebrain activity is altered by sex hormone/receptor changes in males and females (Egozi et al., 1986; Kalinowski et al., 2023; Kirry et al., 2019; Matsuda et al., 2002; Mizuno et al., 2022; van Huizen et al., 1994). For instance, the cholinergic neurons that project to the BLA, originating in the basal forebrain, exhibit high expression of estrogen receptors that is influenced by a female’s hormonal state (Shughrue et al., 2000). During estrus, the enhanced circulating estrogen affects release of ACh and may influence neuronal networks and behavioral phenotypes in a distinct manner (Gibbs, 1996; McEwen, 1998). Thus, increased ACh release in estrus females may underlie increased attentional and risk assessment behaviors in response to vocalization playback, as we observed with flinching behavior in this experimental group. Combined with DA increase, it may trigger both NAc and CeA circuit activation, resulting in both reward-seeking and cautionary behaviors in estrus females.

Revised version: As the results with males listening to restraint vocalizations demonstrate, increased ACh release in BLA is associated with processing aversive cues. How, then, should the increased ACh release in estrus females during mating vocal playback be interpreted? Previous work shows that neuromodulation of amygdalar and other forebrain activity is altered by sex hormone/receptor changes in males and females (Egozi et al., 1986; Kalinowski et al., 2023; Kirry et al., 2019; Matsuda et al., 2002; Mizuno et al., 2022; van Huizen et al., 1994). For instance, the cholinergic neurons that project to the BLA, originating in the basal forebrain, exhibit high expression of estrogen receptors that is influenced by a female’s hormonal state (Shughrue et al., 2000).

Added in the revised version: During estrus, the enhanced circulating estrogen affects release of ACh and may influence neuronal networks and behavioral phenotypes in a distinct manner (Gibbs, 1996; McEwen, 1998). Thus, increased ACh release in estrus females may underlie increased attentional and risk assessment behaviors in response to vocalization playback, as we observed with flinching behavior in this experimental group. Combined with DA increase, it may trigger both NAc and CeA circuit activation, resulting in both reward-seeking and cautionary behaviors in estrus females.

Earlier version: Our results demonstrate the strong impact of even limited experience in shaping behavioral and neuromodulatory responses associated with salient social vocalizations. In the adult mice, a single 90-min session of mating and of restraint, occurring 4-5 days prior to the playback experiment, resulted in consistent behavioral responses and consistent and enhanced ACh and DA release into BLA for both vocalization types. As previous work shows (Huang et al., 2012; Nadim & Bucher, 2014; Pawlak et al., 2010), neuromodulatory inputs play crucial roles in regulating experience-dependent changes in the brain. However, it remains unclear whether the experience shapes neuromodulator release, or whether neuromodulators deliver the experience-related effect into the BLA.

Revised version: Our results demonstrate the strong impact of even limited experience in shaping behavioral and neuromodulatory responses associated with salient social vocalizations. In the adult mice, a single 90-min session of mating and of restraint, occurring 4-5 days prior to the playback experiment, resulted in consistent behavioral responses and consistent and enhanced ACh release into BLA for both vocalization types. As previous work shows (Huang et al., 2012; Nadim & Bucher, 2014; Pawlak et al., 2010), neuromodulatory inputs play crucial roles in regulating experience-dependent changes in the brain. However, it remains unclear whether the experience shapes neuromodulator release, or whether neuromodulators deliver the experience-related effect into the BLA. It is also unknown whether these experiences were specific to the type of vocalization playback associated with the experience, or whether either experience had a more generalized effect on responses to all vocalization playback, as suggested by the broad physiological effects of mating or restraint (Leuner et al 2010; Arnold et al., 2019).

The interaction between ACh and DA is thought to shape motor responses to external stimuli (Lester et al., 2010). Our results support the view that a balance of DA and ACh may regulate the proper behavioral response to appetitive and aversive auditory cues. For instance, increased reward-seeking behavior (rearing and locomotion) in EXP males during mating vocalizations playback may result from the differential release of the two neuromodulators—decreased ACh and increased DA. Further, the lack of this differential release may be the underlying cause for the lack of such responses in INEXP male mice. This supports the role of experience in tuning interactions of these two neuromodulators throughout the BLA, for shaping appropriate behaviors. Overall, the behavioral changes orchestrated by the BLA in response to emotionally salient stimuli are most likely the result of the interaction between previous emotional experiences, hormonal state, content of sensory stimuli, and sex of the listening animals.

Materials and methods

Animals

Experimental procedures were approved by the Institutional Animal Care and Use Committee at Northeast Ohio Medical University. A total of 83 adult CBA/CaJ mice (Jackson Labs, p90-p180), male and female, were used for this study. Animals were maintained on a reversed dark/light cycle and experiments were performed during the dark cycle. The mice were housed in same-sex groups until the week of the experiments, during which they were singly housed. Food and water were provided ad libitum except during the experiment.

The estrous stage of female mice was evaluated based on vaginal smear samples obtained by sterile vaginal lavage. Samples were collected using glass pipettes filled with double distilled water, placed on a slide, stained using crystal violet, and coverslipped for microscopic examination. Estrous stage was determined by the predominant cell type: squamous epithelial cells (estrus), nucleated cornified cells (proestrus), or leukocytes (diestrus) (McLean et al., 2012). To confirm that the stage of estrous did not change during the experiment day, samples obtained prior to and after data collection on the experimental day were compared.

Experimental overview

The basic experimental structure that occurred over six days is illustrated in Figure 1D. Briefly, each mouse was placed in an arena on Days 1 and 2 to provide one-time experiences of mating and sustained restraint. After the experience on Day 2, the subject was anesthetized for implantation of a guide tube for the microdialysis probe. On Day 6, the vocalization playback session occurred. The mouse was briefly anesthetized for insertion of the microdialysis probe into the amygdala, followed by several hours recovery. Before, during, and after vocal playback, we sampled extracellular fluid from the amygdala through the microdialysis probe and recorded video to analyze the subject’s behavior. A liquid chromatographic/mass spectrometric method was used to measure concentrations of several neurochemicals.

The microdialysis sampling interval of 10 mins was used to establish the temporal framework for vocal playback and analysis of neurochemicals and behaviors (Fig. 1E). The analysis compared neurochemicals and behaviors in a period immediately preceding the vocal stimuli (Pre-Stim) and in two 10-min periods featuring playback of the same sequence of vocal stimuli (Stim 1 and Stim 2). Each playback session featured only one category of vocalizations (associated with mating or restraint) and each animal participated in a single playback session. Additional details of this experimental structure are described in the sections below.

Acoustic methods

Vocalization recording and analysis

To record vocalizations for use in playback experiments, mice were placed in an open-topped plexiglass chamber (width, 28 cm; length, 28 cm; height, 20 cm), housed within a darkened, single-walled acoustic chamber (Industrial Acoustics, New York, NY) lined with anechoic foam (Grimsley et al., 2016; Niemczura et al., 2020). Acoustic signals were recorded using ultrasonic condenser microphones (CM16/CMPA, Avisoft Bioacoustics, Berlin, Germany) connected to a multichannel amplifier and A/D converter (UltraSoundGate 416H, Avisoft Bioacoustics). The gain of each microphone was independently adjusted once per recording session to optimize the signal-to-noise ratio (SNR) while avoiding signal clipping. Acoustic signals were digitized at 500 kHz and 16-bit depth, monitored in real time with RECORDER software (Version 5.1, Avisoft Bioacoustics), and Fast Fourier Transformed (FFT) at a resolution of 512 Hz. A night vision camera (VideoSecu Infrared CCTV), centered 50 cm above the floor of the test box, recorded the behaviors synchronized with the vocal recordings (VideoBench software, DataWave Technologies, version 7).

To record mating vocalizations, ten animals (5 male-female pairs) were used in sessions that lasted for 30 minutes. A male mouse was introduced first into the test box, followed by a female mouse 5 min later. Vocalizations were recorded using two ultrasonic microphones placed 30 cm above the floor of the recording box and 13 cm apart. See below for analysis of behaviors during vocal recordings.

To record vocalizations during restraint, six mice (4 male, 2 female) were briefly anesthetized with isoflurane and then placed in a restraint jacket as described previously (Grimsley et al., 2016). Vocalizations were recorded for 30 minutes while the animal was suspended in the recording box. Since these vocalizations are usually emitted at lower intensity compared to mating vocalizations, the recording microphone was positioned 2-3 cm from the snout to obtain the best SNR.

Vocal recordings were analyzed offline using Avisoft-SASLab Pro (version 5.2.12, Avisoft Bioacoustics) with a hamming window, 1024 Hz FFT size, and an overlap percentage of 98.43. For every syllable the channel with the higher amplitude signal was extracted using a custom-written Python code and analyzed. Since automatic syllable tagging did not allow distinguishing some syllable types such as noisy calls and mid-frequency vocalizations (MFVs) from background noise, we manually tagged the start and end of each syllable, then examined spectrograms to measure several acoustic features and classify syllable types based on Grimsley et al. (Grimsley et al., 2011, 2016).

Vocalization playback

Vocalization playback lasting 20 min was constructed from a sequence of seven repeating stimulus blocks lasting 2:50 min each (Fig. 1E). Each block was composed of a set of vocal sequence exemplars that alternated with an equal duration of background sound (“Silence”) associated with the preceding exemplar. The exemplars were recorded during mating interactions and restraint, and selected based on high SNR, correspondence with behavioral category by video analysis, and representation of the spectrotemporal features of vocalizations emitted during mating and restraint (Ghasemahmad, 2020; Grimsley et al., 2016). Mating stimulus blocks contained five exemplars of vocal sequences emitted during mating interactions. These exemplars ranged in duration from 15.0 – 43.6 s. Restraint stimulus blocks included seven vocal sequences, emitted by restrained male or female mice, with durations ranging from 5.7 – 42.3 s. Across exemplars, each stimulus block associated with both mating and restraint included different sets of vocal categories (Fig. 1A-C).

Playback sequences, i.e., exemplars, were conditioned in Adobe Audition CC (2018), adjusted to a 65 dB SNR level, then normalized to 1V peak-to-peak for the highest amplitude syllable in the sequence. This maintained relative syllable emission amplitude in the sequence. For each sequence, an equal duration of background noise (i.e., no vocal or other detected sounds) from the same recording was added at the end of that sequence (Fig. 1E). A 5-ms ramp was added at the beginning and the end of the entire sequence to avoid acoustic artifacts. A MATLAB app (EqualizIR, Sharad Shanbhag; https://github.com/TytoLogy/EqualizIR) compensated and calibrated each vocal sequence for the frequency response of the speaker system. Vocal sequences were converted to analog signals at 500 KHz and 16-bit resolution using DataWave (DataWave SciWorks, Loveland, CO), anti-alias filtered (TDT FT6-2, fc=125KHz), amplified (HCA-800II, Parasound, San Francisco, CA), and sent to the speaker (LCY, K100, Ying Tai Audio Company, Hong Kong). Each sequence was presented at peak level equivalent to 85 dB SPL.

Behavioral methods

Behaviors during both vocalization recording and playback sessions were recorded using a night vision camera (480TVL 3.6mm, VideoSecu), centered 50 cm above the floor of the test box, and SciWorks (DataWave, VideoBench version 7) for video acquisition and analysis.

Analysis of mating behaviors during vocal recordings

Interactions between male and female mice were video-recorded, then analyzed second-by-second. Among more general courtship interactions, we identified a set of behaviors that are mating-related, as described previously (Gaub et al., 2016; Heckman et al., 2016). All vocal sequences selected as exemplars for playback of “mating” vocalizations were recorded in association with these male mating behaviors: head-sniffing, attempted mounting, or mounting. Vocalizations during these behaviors included chevron, stepped, and complex USVs emitted with longer durations and higher repetition rates, and more LFH calls (Gaub et al., 2016; Ghasemahmad, 2020; Hanson & Hurley, 2012).

Experience and playback sessions

Prior to playback experiments, each animal underwent 90-min sessions on two consecutive days (Days 1 and 2) that provided both mating and restraint experiences to the EXP group (n = 31 animals) or no experiences of these to the INEXP group (n = 22 animals). EXP sessions were presented in a counterbalanced pattern across subjects (Fig. 1D). For the mating experience, mounting or attempted mounting was required for the animal to be included in the remainder of the experiment. However, we did not record detailed behaviors or track estrous stage during the mating experience session. After the Day 2 session, mice underwent surgery for implantation of a microdialysis guide canula (see below), then recovered for four days.

On Day 6, the day of the playback experiment, male mice were randomly assigned to either restraint or mating vocal playback groups. Females were only tested with mating playback, since a preliminary behavioral study suggested no difference in male and female responses to restraint. Video recording was performed simultaneously with microdialysis, beginning 10 min before vocal playback (pre-stimulus) and continuing for 20 min during playback (Fig. 1D).

Analysis of behaviors during vocal playback

Behavioral analysis was based on previous descriptions (Bakshi & Kelley, 1993, 1994; Blanchard et al., 2003; Füzesi et al., 2016; Grimsley et al., 2015; Lezak et al., 2017; Saibaba et al., 1996) (Mouse Ethogram: https://mousebehavior.org). The list of behaviors assessed was determined through pilot analysis and previous studies characterizing rodent defense and fear behaviors. They are defined as the following: abrupt attending, flinching, locomotion, rearing, grooming, still-and-alert, and stretch-attend (description in Table 1; (Blanchard et al., 2003)). A behavior was only counted when its occurrence lasted two or more seconds, except for flinching, which could take place more quickly.

For analysis of behaviors during vocal playback, all videos were examined blind to the sex of the animal and the context of vocalizations. Video recordings were analyzed manually (n=48 mice) in 10-second intervals for the seven behaviors identified above. The number of occurrences of every behavior per 10-second block was marked. The numbers were added for every block of 2:50 min of vocal playback (see Fig. 1E), then averaged for the Pre-Stim, Stim 1, and Stim 2 periods separately.

Videos were further analyzed automatically using the video tracker within VideoBench (DataWave Technologies, version 7) for total distance traveled and time spent in the periphery and center of the test arena. For video-tracking analysis of time spent, the floor was divided into 16 equal squares with 4 central squares identified as “center” and the 4 corner squares identified as periphery. The two squares in the middle of each wall were excluded from the analysis to avoid attraction to the holes affect the animal’s preference for the periphery vs center.

Procedures related to microdialysis

Surgery

Earlier version: Since both left and right amygdala are responsive to vocal stimuli in human and experimental animal studies (Wenstrup et al., 2020), we implanted microdialysis probes into the left amygdala to maintain consistency with other studies in our laboratory. Mice were anaesthetized with Isoflurane (2-4%, Abbott Laboratories, North Chicago, IL) and hair overlying the skull was removed using depilatory lotion. A midline incision was made, then the skin was moved laterally to expose the skull. A craniotomy (≃1 mm2) was made above the basolateral amygdala (BLA) on the left side (stereotaxic coordinates from bregma: −1.65 mm rostrocaudal, +3.43 mm mediolateral). A guide cannula (CMA-7, CMA Microdialysis, Sweden) was implanted to a depth of 2.6 mm below the cortical surface and above the left BLA (Fig. 1D, Day 2), then secured using dental cement. After surgery, the animal received a subcutaneous injection of carprofen (4 mg/kg, s.c.) and topical anesthetic (lidocaine) and antibiotic cream (Neosporin). It was returned to its cage and placed on a heating pad until full recovery from anesthesia. Recovery from the surgical procedures occurred in the animal’s home cage and was assessed by resumption of normal eating, drinking, exploratory, and grooming behaviors and postures. Playback experiments occurred in all recovered animals four days after surgery.

Revised version: Since both left and right amygdala are responsive to vocal stimuli in human and experimental animal studies (Wenstrup et al., 2020), we implanted microdialysis probes into the left amygdala to maintain consistency with other studies in our laboratory. Mice were anaesthetized with Isoflurane (2-4%, Abbott Laboratories, North Chicago, IL) and hair overlying the skull was removed using depilatory lotion. A midline incision was made, then the skin was moved laterally to expose the skull. A craniotomy (≃1 mm2) was made above the basolateral amygdala (BLA) on the left side (stereotaxic coordinates from bregma: -1.65 mm rostrocaudal, +3.43 mm mediolateral). A guide cannula (CMA-7, CMA Microdialysis, Sweden) was implanted to a depth of 2.6 mm below the cortical surface and above the left BLA (Fig. 1D, Day 2), then secured using dental cement. After surgery, the animal received a subcutaneous injection of carprofen (4 mg/kg, s.c.) and topical anesthetic (lidocaine) and antibiotic cream (Neosporin). It was returned to its cage and placed on a heating pad until fully recovered from anesthesia. Recovery from the surgical procedures occurred in the animal’s home cage and was assessed by resumption of normal eating, drinking, exploratory, and grooming behaviors and postures. Playback experiments occurred in all recovered animals four days after surgery.

Microdialysis

On the day before microdialysis, the probe was conditioned in 70% methanol and artificial cerebrospinal fluid (aCSF) (CMA Microdialysis, Sweden). On playback day (Day 6), the animal was briefly anesthetized and the probe with 1 mm membrane length and 0.24 mm outer diameter (MWCO 6 kDa) was inserted into the guide cannula (Fig. 1D).

Using a spiral tubing connector (0.1 mm ID x 50 cm length) (CT-20, AMUZA Microdialysis, Japan), the inlet and outlet tubing of the probe was connected to the inlet /outlet Teflon tubing of the microdialysis lines. A swivel device for fluids (TCS-2-23, AMUZA Microdialysis, Japan), secured to a balance arm, held the tubing, and facilitated the animal’s free movement during the experiment.

Earlier version: After probe insertion, a four-hour period allowed animals to habituate and the neurochemicals to equilibrate between aCSF fluid in the probe and brain extracellular fluid (ECF). Sample collection then began, in 10-min intervals, beginning with four background samples, then two samples during playback of restraint or mating vocal sequences (total 20 minutes), then one or more samples after playback ended. To account for the dead volume of the outlet tubing, a flow rate of 1.069 µl/min was established at the syringe pump to obtain a 1 µl sample per minute. Samples collected during this time were measured for volume to assure a consistent flow rate. To prevent degradation of collected neurochemicals, the outlet tubing was passed through ice to a site outside the sound-proof booth where samples were collected on ice, then stored in a −80°C freezer. The choice of 10-min sampling intervals was based on the minimum collection time needed to provide 5µl samples both for immediate testing and for a backup sample in the event of catastrophic failure of the processes associated with the neurochemical analysis.

Revised version: After probe insertion, a four-hour period allowed animals to habituate and the neurochemicals to equilibrate between aCSF fluid in the probe and brain extracellular fluid (ECF). Sample collection then began, in 10-min intervals, beginning with four background samples, then two samples during playback of restraint or mating vocal sequences (total 20 minutes), then one or more samples after playback ended. To account for the dead volume of the outlet tubing, a flow rate of 1.069 µl/min was established at the syringe pump to obtain a 1 µl sample per minute.

Added in the revised version: Samples collected during this time were measured for volume to assure a consistent flow rate. To prevent degradation of collected neurochemicals, the outlet tubing was passed through ice to a site outside the sound-proof booth where samples were collected on ice, then stored in a −80°C freezer. The choice of 10-min sampling intervals was based on the minimum collection time needed to provide 5µl samples both for immediate testing and for a backup sample in the event of catastrophic failure of the processes associated with the neurochemical analysis.

Neurochemical analysis

Samples were analyzed using a liquid chromatography (LC) / mass spectrometry (MS) technique (Fig. 1D) at the Vanderbilt University Neurochemistry Core. This method allows simultaneous measurement of several neurochemicals in the same dialysate sample: acetylcholine (ACh), dopamine (DA) and its metabolites (3,4-Dihydroxyphenylacetic acid (DOPAC), and homovanillic acid (HVA)), serotonin (5-HT) and its metabolite (5-hydroxyindoleacetic acid (5-HIAA)), norepinephrine (NE), gamma aminobutyric acid (GABA), and glutamate. Due to low recovery of NE and 5-HT from the mouse brain, we were unable to track these two neuromodulators in this experiment.

Before each LC/MS analysis, 5 μl of the sample was derivatized using sodium carbonate, benzoyl chloride in acetonitrile, and internal standard (Wong et al., 2016). LC was performed on a 2.0 × 50 mm, 1.7 μM particle Acquity BEH C18 column (Waters Corporation, Milford, MA, USA) with 1.5% aqueous formic acid as mobile phase A and acetonitrile as mobile phase B. Using a Waters Acquity Classic UPLC, samples were separated by a gradient of 98–5 % of mobile phase A over 11 min at a flow rate of 0.6 ml/min with delivery to a SCIEX 6500+ Qtrap mass spectrometer (AB Sciex LLC, Framingham, MA, USA). The mass spectrometer was operated using electrospray ionization in the positive ion mode. The capillary voltage and temperature were 4KV and 350°C, respectively (Wong et al., 2016; Yohn et al., 2020). Chromatograms were analyzed using MultiQuant 3.0.2 Software (AB SCIEX, Concord, Ontario, Canada).

Verification of probe location

We verified placement of microdialysis probes to minimize variability that could arise due to placement in regions surrounding BLA that receive different sources of neurochemical inputs (e.g., cholinergic inputs to putamen and central amygdala). To verify probe placement within the BLA, each probe was perfused with 2% dextran-fluorescein (MW 4Kda) (Sigma Aldrich Inc, Atlanta, GA) at the end of the experiment. The location of the probe was then visualized in adjacent cleared and Nissl-stained sections (e.g., Fig. 2, insets). Sections were photographed using a SPOT RT3 camera and SPOT Advanced Plus imaging software (version 4.7) mounted on a Zeiss Axio Imager M2 fluorescence microscope. Adobe Photoshop CS3 was used to adjust brightness and contrast globally. Animals were included in statistical analyses only if ≥75% of the probe membrane was located within the BLA (Fig. 2).

Data analysis

Removed in the revised version: For neurochemical analyses, the total numbers of animals used were NEXP=31 and NINEXP=22. Only mice with useable neurochemical data were further evaluated for behaviors. This included 46 mice: (NEXP=25 and NINEXP=21). Note that some animals used in neurochemical analyses were not included in the behavioral analyses because their behavioral data were unavailable. Since only one INEXP female was in a non-estrus stage during the playback session, our analysis of the effect of experience included only estrus females and males. Further, one INEXP female without detectable DA levels was removed from the INEXP group neurochemical analysis.

Earlier version: Only the following 10-min sample collection windows were used for statistical analysis: Pre-Stim, Stim 1, and Stim 2 (Fig. 1E). The purpose of the two collection windows during vocal playback was to allow detection of different epochs of neurochemical release, as may occur in ACh release into the amygdala (Aitta-aho et al., 2018). All neurochemical data were normalized to the background level, obtained from a single pre-stimulus sample immediately preceding playback. This provided clarity in representations and did not result in different outcomes of statistical tests compared to use of three pre-stimulus samples. The fluctuation in all background samples was ≤20%. The percentage change from background level was calculated based on the formula: % change from background = (100 x stimulus sample concentration in pg)/ background sample concentration in pg. Values are represented as mean ± one standard error unless stated otherwise. Raw neurochemical values are reported in supplementary tables 1 (ACh), 2 (DA), and 3 (5-HIAA).

Revised version: For neurochemical analyses, the total numbers of animals used were NEXP=31 and NINEXP=22. Only mice with usable neurochemical data were further evaluated for behaviors. This included 46 mice: (NEXP=25 and NINEXP=21). Note that some animals used in neurochemical analyses were not included in the behavioral analyses because their behavioral data were unavailable. Since only one INEXP female was in a non-estrus stage during the playback session, our analysis of the effect of experience included only estrus females and males. Further, one INEXP female without detectable DA levels was removed from the INEXP group neurochemical analysis. Only the following 10-min sample collection windows were used for statistical analysis: Pre-Stim, Stim 1, and Stim 2 (Fig. 1E). The purpose of the two collection windows during vocal playback was to allow detection of different epochs of neurochemical release, as may occur in ACh release into the amygdala (Aitta-aho et al., 2018). All neurochemical data were normalized to the background level, obtained from a single pre-stimulus sample immediately preceding playback. This provided clarity in representations and did not result in different outcomes of statistical tests compared to use of three pre-stimulus samples. The fluctuation in all background samples was ≤20%. The percentage change from background level was calculated based on the formula: % change from background = (100 x stimulus sample concentration in pg)/ background sample concentration in pg. Values are represented as mean ± one standard error unless stated otherwise. Raw neurochemical values are reported in supplementary tables 1 (ACh), 2 (DA), and 3 (5-HIAA).

All statistical analyses were performed using SPSS (IBM, V. 26 and 27). To examine vocalizations in different stages of mating, we used a linear mixed model to analyze the changes of interval and duration (dependent variables) of vocalizations based on mating interaction intensity (fixed effect). For behavioral and neurochemical analyses in playback experiments, we initially compared the output using a linear mixed model and a generalized linear model (GLM) with a repeated measure for neurochemical data. Since both statistical methods resulted in similar findings, we chose to use the GLM for statistical comparisons of the neurochemicals and behaviors. Where Mauchly’s test indicated that the assumption of sphericity had been violated, degrees of freedom were corrected using Greenhouse Geisser estimates of sphericity.

To further clarify the differences observed between groups for every comparison in the GLM repeated measure, multivariate contrast was performed. All multiple comparisons were corrected using Bonferroni post hoc testing. 95% confidence intervals were used to compare values between timepoints (Julious, 2004).

Earlier version: For both microdialysis and behavioral data, we tested the hypotheses that the release of neurochemicals into the BLA and the number of behaviors is differently modulated by the vocal playback type (mating and restraint) in male mice, by estrous stage of females or sex of the animals during mating vocal playback, and by previous experience in any of these groups. The GLM model for testing these hypotheses used time as a within-subject factor, with vocalization context, sex, estrous, or the interaction of these with EXP and time as the between-subject factors. Dependent (response) variables included the normalized concentration of ACh, DA, 5-HIAA or the numbers of behaviors as previously described in the behavioral analysis section.

Revised version: For both microdialysis and behavioral data, we tested the hypotheses that the release of neurochemicals into the BLA and the number of behaviors is differently modulated by the vocal playback type (mating and restraint) in male mice, by estrous stage of females or sex of the animals during mating vocal playback, and by previous experience in any of these groups. In the GLM model for Figure 3, context of vocalizations was defined as the fixed factor and time (Stim 1 and Stim 2) as the within-subject factor. The model in Figure 4 tested the effect of sex and estrous state as a nested variable within sex as fixed factors, which were considered in the model design (design = intercept + Sex + Estrous(sex)). The GLM model in Figure 5 tested sex, context, and experience as fixed factors with the interaction terms of sex*experience and context*experience. Here we avoided using a full factorial model due to the absence of females in the restraint group, as that would have made the model unbalanced (design = intercept + Sex + Context + Experience + Sex*Experience + Context*Experience). Further, estrous state was not tested in examining the experience effect due to lack of INEXP non-estrus females. Dependent (response) variables included the normalized concentration of ACh, DA, 5-HIAA in Stim 1 and Stim 2 or the numbers of behaviors in Pre-Stim, Stim 1, and Stim 2 as previously described in the behavioral analysis section.

Throughout the figures, the distributions of data points are visualized by box plots, computed using the “inclusive median” formula in Excel. The box indicates first (Q1) and third (Q3) quartiles with the median line between the two. The whiskers normally stretch between minimum and maximum values, but this computation sometimes shows data points outside of the whiskers. Excel identifies these values as outliers in the boxplot function. However, our GLM statistical analysis includes these values because there was no compelling reason to exclude them. The mean value is marked with an “x”.

Supporting information